Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Exercise 1 
(10 marks)

$\newcommand{\bfx}{\mathbf{x}}
\newcommand{\bfb}{\mathbf{b}}$
Find the minimum of the following function
$$f(\bfx) = \bfx^\top A^\top A \bfx + 2 \bfb^\top \bfx + c$$


## Exercise 2
(10 marks)

$\newcommand{\bfmu}{\boldsymbol{\mu}}$
Find the minimum of the following function (assume $\Sigma_1$ and $\Sigma_2$ to be positive semi-definite).
$$f(\bfx) = (\bfx - \bfmu_1)^\top \Sigma_1 (\bfx - \bfmu_1)
+ (\bfx - \bfmu_2)^\top \Sigma_2 (\bfx - \bfmu_2) $$

## Exercise 3: Fit a plane to point cloud

Least square fit in 2D fits a line to a given set of 2D points. This homework asks you to repeat the minimization procedure to find a plane that fits a 2D points. Below you are provided with helper code to load a point cloud, and visualize the points and the plane. Your task is to find the equation of plane that fits the point cloud. 

1. Show the math and derivation for your code
2. Implement the code in python and numpy to find the coefficients of the equation of plane and write your answer in the space provided below. Submit this ipynb file to brightspace. (30 marks)

```python
#Comment after running this once
!pip install --user k3d


import k3d
from k3d import matplotlib_color_maps

def k3d_init_plot():
    return k3d.plot(grid_visible=False,
                    camera_auto_fit=True,
                    camera_fov=30.0,
                    background_color=0xffffff)

def visualize_plane(pts, abc, k3dplt=None):
    _k3dplt = k3dplt
    k3dplt = k3d_init_plot() if k3dplt is None else k3dplt
    a, b, c = abc
    x = np.linspace(-1, 1, 100, dtype=np.float32)
    y = np.linspace(0, 1, 100, dtype=np.float32)

    x, y = np.meshgrid(x, y)
    z = a*x + b*y + c
    plt_surface = k3d.surface(z,
                              xmin=-1, xmax=1,
                              ymin=0, ymax=1,
                              compression_level=9,
                              wireframe=True,
                              color_map=matplotlib_color_maps.Coolwarm_r,
                              attribute=z, color_range=[-1, 0.5])
    
    k3dplt += plt_surface
    visualize_points(pts, k3dplt=k3dplt)
    if _k3dplt is None:
        k3dplt.display()

def visualize_points(pts, k3dplt=None):
    _k3dplt = k3dplt
    k3dplt = k3d_init_plot() if k3dplt is None else k3dplt
    plt_points = k3d.points(pts,
                            attribute=pts[:, 2], # use z coordinate to color
                            point_size=0.005,
                            shader="flat",
                            color_map=matplotlib_color_maps.Inferno)
    k3dplt += plt_points
    if _k3dplt is None:
        k3dplt.display()
```

```python
import numpy as np
!wget https://github.com/wecacuee/ECE490-Neural-Networks/raw/master/notebooks/02-linear-models/table_top.npy
pts = np.load('table_top.npy').astype('float32') # Nx3
visualize_points(pts)
pts.min(axis=0), pts.max(axis=0)
```

<!-- grade_id: cell-bb7ef5a91325977a, points: 0, grade: false, solution: true -->
```python
def plane_fit(pts):
    # YOUR CODE HERE
    raise NotImplementedError()
    return abc
```

<!-- grade_id: cell-5b7358c1e61a4706, points: 50, grade: true, solution: false -->
```python
# Find abc such that z = a x + b y + c
abc = plane_fit(pts)
visualize_plane(pts, abc)
assert np.abs(pts[:,2] - np.hstack((pts[:, :2], np.ones((pts.shape[0],1)))) @ abc).mean() < 0.4

```